# Lab 08: Functions and Tables


**Helpful Resource:**
- [Python Reference](http://data8.org/sp24/reference/): Cheat sheet of helpful array & table methods!

**Recommended Readings**:

* [Visualizing Numerical Distributions](https://www.inferentialthinking.com/chapters/07/2/Visualizing_Numerical_Distributions.html)
* [Functions and Tables](https://www.inferentialthinking.com/chapters/08/Functions_and_Tables.html)

Please complete this notebook by filling in the cells provided. **Before you begin, execute the cell below to setup the notebook by importing some helpful libraries.** Each time you start your server, you will need to execute this cell again.

```python
!apt-get install texlive texlive-xetex texlive-latex-extra pandoc
!pip install pypandoc
!pip install datascience
```

```python
# Connect Google Drive to Colab so you can access your files
from google.colab import drive
drive.mount('/content/drive')

import os
os.chdir('/content/drive/MyDrive/Colab Notebooks/')
```

```python
# Run this cell to set up the notebook, but please don't change it.

# These lines import the Numpy and Datascience modules.
import numpy as np
from datascience import *

# These lines do some fancy plotting magic.
import matplotlib
%matplotlib inline
import matplotlib.pyplot as plt
plt.style.use('fivethirtyeight')

import warnings
warnings.simplefilter('ignore', FutureWarning)
warnings.filterwarnings("ignore")
```

## 1. Distribution of Durations



```python
trip = Table.read_table('./DS/trip.csv')
trip.show(3)
```

```python
trip.hist('Duration')
```

```python
trip.sort('Duration', descending=True)
```

<br>

**Question 1.** Create histogram with bins ranging from 0 to 1800 in increments of 250

```python
trip.hist('Duration', bins=np.arange(0, 1800, 250))
```

<br>

**Question 2.** Create histogram with more fine grained bins

```python
trip.hist('Duration', bins=np.arange(0, 1800))
```

## 2. Burrito-ful San Diego



Edwin, Gamy and Andrew are trying to use Data Science to find the best burritos in San Diego! Their friends Jessica, Aileen, and Conan provided them with two comprehensive datasets on many burrito establishments in the San Diego area taken and cleaned from [here](https://www.kaggle.com/srcole/burritos-in-san-diego/data)

The following cell reads in a table called `ratings` which contains names of burrito restaurants, their Yelp rating, Google rating, as well as their overall rating. The `Overall` rating is not an average of the `Yelp` and `Google` ratings, but rather it is the overall rating of the customers that were surveyed in the study above.


It also reads in a table called `burritos_types` which contains names of burrito restaurants, their menu items, and the cost of the respective menu item at the restaurant.

```python
# Just run this cell
ratings = Table.read_table("./DS/ratings.csv")
ratings.show(5)
burritos_types = Table.read_table("./DS/burritos_types.csv").drop(0)
burritos_types.show(5)
```

**Question 1.** It would be easier if we could combine the information in both tables. Assign `burritos` to the result of joining the two tables together, so that we have a table with the ratings for every corresponding menu item from every restaurant. Each menu item has the same rating as the restaurant from which it is from. **(8 Points)**

*Hint:* Which function lets us combine information from two tables? Refer to the [Python Reference Sheet](http://data8.org/sp24/reference/) if you're unsure!


```python
burritos = ratings.join('Name', burritos_types)
burritos.show(10)
```

<!-- BEGIN QUESTION -->

**Question 2.** Let's look at how the Yelp scores compare to the Google scores in the `burritos` table. First, assign `yelp_google_tbl` to a table only containing the columns `Yelp` and `Google`. Then, make a scatter plot with Yelp scores on the x-axis and the Google scores on the y-axis. **(8 Points)**


```python
yelp_google_tbl = burritos.select('Yelp', 'Google')
yelp_google_tbl.scatter('Yelp', 'Google')

# Don't change/edit/remove the following line.
# To help you make conclusions, we have plotted a straight line on the graph (y=x).
plt.plot(np.arange(2.5,5,.5), np.arange(2.5,5,.5));
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 3.** Looking at the scatter plot you just made in Question 1.2, do you notice any pattern(s) or relationships between Yelp and Google ratings (i.e. is one of the two types of scores consistently higher than the other one)? If so, describe them **briefly** in the cell below. **(8 Points)**


Based on the scatter plot, we can see that there are more Google points than Yelp points, especially at scores above 4.0. Although there is a tendency for one rating to increase as the other increases, the high values of the Google ratings are mostly consistently above the y=x line.

<!-- END QUESTION -->

Here's a refresher on how `.group` works! You can read how `.group` works in the [textbook](https://www.inferentialthinking.com/chapters/08/2/Classifying_by_One_Variable.html).

**Question 4.** There are so many types of California burritos in the `burritos` table! Gamy wants to consider his options for burritos based on rankings. For the sake of these questions, we are treating each menu item's rating the same as its respective restaurant's, as we do not have the rating of every single item at these restaurants. You do not need to worry about this fact, but we thought to mention it!

Create a table with two columns: the first column include the **names of the burritos** and the second column should contain the **average overall rating** of that burrito across restaurants. **In your calculations, you should only compare burritos that contain the word "California".** For example, there are "California" burritos, "California Breakfast" burritos, "California Surf And Turf" burritos, etc. **(9 Points)**

*Hint:* “California” is case sensitive (“California” is not the same as “california”)!

*Note:* For reference, the staff solution only used one line. However, feel free to break up the solution into multiple lines and steps; just make sure you assign the final output table to `california_burritos`!


```python
california_burritos = burritos.where("Menu_Item", are.containing("California")).group("Menu_Item", np.average).select("Menu_Item", "Overall average")
california_burritos
```

**Question 5.** Given this new table `california_burritos`, Andrew can figure out the name of the California burrito with the highest overall average rating! Assign `best_california_burrito` to a line of code that outputs the string that represents the name of the California burrito with the highest overall average rating. If multiple burritos satisfy this criteria, you can output any one of them. **(8 Points)**


```python
best_california_burrito = california_burritos.sort("Overall average", descending=True).column("Menu_Item").item(0)
best_california_burrito
```

<!-- BEGIN QUESTION -->

**Question 6.** Edwin thinks that burritos in San Diego are cheaper (and taste better) than the burritos in Berkeley. Plot a histogram that visualizes that distribution of the costs of the burritos from San Diego in the `burritos` table. Also use the provided `cost_bins` variable when making your histogram, so that the histogram is more visually informative. **(8 Points)**


```python
cost_bins = np.arange(0, 15, 1) # Do not change this line
# Please also use the provided bins
burritos.hist("Cost", bins=cost_bins)
```

<!-- END QUESTION -->

**Question 7.** What percentage of burritos in San Diego are less than $6? Assign `burritos_less_than_six` to your answer, **which should be between 0 and 100**. You should only use the histogram above to answer the question. Do not use code on the table to to find the answer, just eyeball the heights and use Python to evaluate your arithmetic! **(8 Points)**

*Note*: Your answer does not have to be exact, but it should be within a couple percentages of the staff answer.



```python
burritos_less_than_six = 1 + 4 + 17
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a pdf file for you to submit. **Please save before exporting!**

```python
# should change the directory and file name matching to yours
!jupyter nbconvert './DS/lab08(SSU).ipynb' --to pdf
```

Tests for q1_1 (points: [0, 0])
```
>>> # Make sure you have all the columns from both tables
>>> set(["Name", "Menu_Item", "Yelp", "Google", "Overall", "Cost"]) == set(burritos.labels)
True
```
```
>>> burritos.num_rows == 212
True
```


Tests for q1_4 (points: [1, 1])
```
>>> # Number of columns should be 2
>>> california_burritos.num_columns == 2
True
```
```
>>> # Number of rows should be 19
>>> california_burritos.num_rows == 19
True
```


Tests for q1_5 (points: [0])
```
>>> # Your answer should be a string
>>> type(best_california_burrito) == str
True
```


Tests for q1_7 (points: [0])
```
>>> # Your answer should be between 0 and 100.
>>> 0 <= burritos_less_than_six <= 100
True
```


Tests for q2_1 (points: [0, 0, 0, 0])
```
>>> job_titles.num_columns
2
```
```
>>> job_titles.num_rows
6
```
```
>>> # Make sure that you have the correct column labels!
>>> np.asarray(job_titles.labels).item(1) != "Job full_array"
True
```
```
>>> # Make sure that you have the correct column labels!
>>> np.asarray(job_titles.labels).item(1) == "Jobs"
True
```


Tests for q2_3 (points: [0, 0, 2])
```
>>> # Double check that your salary_range function is correct
>>> compensation_range(make_array(5, 1, 20, 1000)) == 999
True
```
```
>>> # Make sure that the table has departments as the rows and organization groups as the columns.
>>> set(["Department", "Community Health", "Culture & Recreation", "General Administration & Finance", "Human Welfare & Neighborhood Development", "Public Protection", "Public Works, Transportation & Commerce"]) == set(department_ranges.labels)
True
```
```
>>> sum(department_ranges.column(1))
554179.0
```


Tests for q2_5 (points: [2, 2])
```
>>> # Make sure that your final answer is a number
>>> isinstance(num_over_125k, int)
True
```
```
>>> # Make sure that your answer makes sense given the sf table
>>> 0 <= num_over_125k <= 51
True
```
