
#### **Helpful Resource:**
- [Python Reference](http://data8.org/sp24/reference/)

**Recommended Readings:**

* [Variability](https://inferentialthinking.com/chapters/14/2/Variability.html#)
* [Correlation](https://inferentialthinking.com/chapters/15/1/Correlation.html#)

```python
!apt-get install texlive texlive-xetex texlive-latex-extra pandoc
!pip install pypandoc
!pip install datascience
```

```python
# Connect Google Drive to Colab so you can access your files
from google.colab import drive
drive.mount('/content/drive')

import os
os.chdir('/content/drive/MyDrive/Colab Notebooks/')
```

```python
# Run this cell, but please don't change it.

# These lines import the Numpy and Datascience modules.
import numpy as np
from datascience import *

# These lines do some fancy plotting magic.
import matplotlib
%matplotlib inline
import matplotlib.pyplot as plots
plots.style.use('fivethirtyeight')
import warnings
warnings.simplefilter('ignore', FutureWarning)
```

# 1. How Faithful is Old Faithful?

Old Faithful is a geyser in Yellowstone National Park that is famous for eruption on a fairly regular schedule.

```python
# For the curious: this is how to display a YouTube video in a
# Jupyter notebook.  The argument to YouTubeVideo is the part
# of the URL (called a "query parameter") that identifies the
# video.  For example, the full URL for this video is:
#   https://www.youtube.com/watch?v=wE8NDuzt8eg
from IPython.display import YouTubeVideo
YouTubeVideo("wE8NDuzt8eg")
```

Some of Old Faithful's eruptions last longer than others.  Whenever there is a long eruption, it is usually followed by an even longer wait before the next eruption. If you visit Yellowstone, you might want to predict when the next eruption will happen, so that you can see the rest of the park instead of waiting by the geyser.

Today, we will use a dataset on eruption durations and waiting times to see if we can make such predictions accurately with linear regression.

The dataset has one row for each observed eruption.  It includes the following columns:
- `duration`: Eruption duration, in minutes
- `wait`: Time between this eruption and the next, also in minutes

Run the next cell to load the dataset.

```python
faithful = Table.read_table("./DS/faithful.csv")
faithful
```

**Question 1.0.** The following statements are the unordered steps of linear regression.  

1. Compute the parameters of the regression line: the slope and the intercept.
2. Evaluate the regression line by computing the line’s RMSE and analyzing the residuals plot.
3. Use the regression line to generate predictions for each x value.
4. Determine if linear regression is a reasonable method by visualizing your data and computing the correlation coefficient.

Make an array called `least_squares_order` that contains the correct order of a linear regression analysis, where the first item of the array is the first step of an linear regression analysis and the last item of the array is the last step of an linear regression analysis.


```python
least_squares_order = make_array(4, 1, 3, 2)
least_squares_order
```

We would like to use linear regression to make predictions, but that won't work well if the data aren't roughly linearly related.  To check that, we should look at the data.

**Question 1.1.** Make a scatter plot of the data.  It's conventional to put the column we want to predict on the vertical axis and the other column on the horizontal axis.


```python
faithful.scatter('duration', 'wait')
```

**Question 1.2.** Are eruption duration and waiting time roughly linearly related based on the scatter plot above? Is this relationship positive?


The scatter plot shows a similar relationship between eruption time and waiting time, and this relationship is positive. That is, the timing of an eruption tends to adjust the waiting time until the next eruption.

We're going to continue with the assumption that they are linearly related, so it's reasonable to use linear regression to analyze this data.

We'd next like to plot the data in standard units. If you don't remember the definition of standard units, textbook section [14.2](https://www.inferentialthinking.com/chapters/14/2/Variability.html#standard-units) might help!

**Question 1.3.** Compute the mean and standard deviation of the eruption durations and waiting times.  **Then** create a table called `faithful_standard` containing the eruption durations and waiting times in standard units.  The columns should be named `duration (standard units)` and `wait (standard units)`.


```python
duration_mean = faithful.column('duration').mean()
duration_std = faithful.column('duration').std()
wait_mean = faithful.column('wait').mean()
wait_std = faithful.column('wait').std()

faithful_standard = Table().with_columns(
    "duration (standard units)", (faithful.column('duration') - duration_mean) / duration_std,
    "wait (standard units)", (faithful.column('wait') - wait_mean) / wait_std)
faithful_standard
```

**Question 1.4.** Plot the data again, but this time in standard units.


```python
faithful_standard.scatter('duration (standard units)', 'wait (standard units)')
```

You'll notice that this plot looks the same as the last one!  However, the data and axes are scaled differently.  So it's important to read the ticks on the axes.

**Question 1.5.** Among the following numbers, which would you guess is closest to the correlation between eruption duration and waiting time in this dataset?

1. -1
2. 0
3. 1

Assign `correlation` to the number corresponding to your guess (either 1, 2 or 3).


```python
correlation = 3
correlation
```

**Question 1.6.** Compute the correlation coefficient: `r`.  

*Hint:* Use `faithful_standard`.  Section [15.1](https://www.inferentialthinking.com/chapters/15/1/Correlation.html#calculating-r) explains how to do this.



```python
r = (faithful_standard.column('duration (standard units)') * faithful_standard.column('wait (standard units)')).mean()
r
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a pdf file for you to submit. **Please save before exporting!**

```python
# should change the directory and file name matching to yours
!jupyter nbconvert './DS/lab23(SSU).ipynb' --to pdf
```

Tests for q1_0 (points: None)
```
>>> len(least_squares_order) == 4
True
```
```
>>> 
>>> import hashlib # This imports a hashing library for the autograder.
>>> def get_hash(num):
...     """Helper function for assessing correctness."""
...     return hashlib.md5(str(num).encode()).hexdigest()
>>> 
>>> get_hash(np.array(least_squares_order).astype(int)) # Your ordering is incorrect.
'cd29f2d730e11535cef30ecc78640daa'
```


Tests for q1_3 (points: None)
```
>>> abs(sum(faithful_standard.column(0))) <= 1e-8
True
```
```
>>> round(duration_std, 2) == 1.14
True
```
```
>>> round(wait_std, 2) == 13.57 
True
```


Tests for q1_5 (points: None)
```
>>> 
>>> import hashlib # This imports a hashing library for the autograder.
>>> def get_hash(num):
...     """Helper function for assessing correctness."""
...     return hashlib.md5(str(num).encode()).hexdigest()
>>> 
>>> get_hash(int(correlation)) # Your choice is incorrect.
'eccbc87e4b5ce2fe28308fd9f2a7baf3'
```


Tests for q1_6 (points: None)
```
>>> -1 <= r <= 1
True
```
```
>>> round(r,3) == 0.901
True
```


Tests for q2_1 (points: None)
```
>>> (slope*13 - 100)/98 <= 0.5
True
```
```
>>> np.round(slope, 4) == 10.7296 # Make sure you are plugging in the SD_x and SD_y in the correct spots!
True
```


Tests for q2_2 (points: None)
```
>>> 33 < intercept < 34
True
```
```
>>> np.round(intercept, 3) == 33.474
True
```


Tests for q3_1 (points: None)
```
>>> round(two_minute_predicted_waiting_time, 3) == 54.934
True
```
```
>>> round(five_minute_predicted_waiting_time, 3) == 87.123
True
```


Tests for q3_2 (points: None)
```
>>> set(faithful_predictions.labels) == set(['duration', 'wait', 'predicted wait']) # Make sure your column labels are correct.
True
```
```
>>> abs(1 - np.mean(faithful_predictions.column(2))/100) <= 0.35
True
```


Tests for q3_3 (points: None)
```
>>> set(faithful_residuals.labels) == set(['duration', 'wait', 'predicted wait', 'residual']) # Make sure your column labels are correct.
True
```
```
>>> abs(sum(faithful_residuals.column('residual'))) <= 1e-8
True
```


Tests for q4_1 (points: None)
```
>>> 12 - zero_minute_predicted_waiting_time*1.4/4 <= 0.35
True
```
```
>>> zero_minute_predicted_waiting_time - intercept <= 1e5
True
```
```
>>> 2 - two_point_five_minute_predicted_waiting_time/35 <= 0.4
True
```
```
>>> (26 - hour_predicted_waiting_time/30)/10 <= 0.43
True
```


Tests for q5_1 (points: None)
```
>>> [below_3.num_rows, above_3.num_rows]
[97, 175]
```
```
>>> np.allclose([below_3_r, above_3_r], [0.290189526493, 0.372782225571])
True
```


Tests for q5_2 (points: None)
```
>>> np.allclose([wait_below_3(1), wait_below_3(3), wait_above_3(3), wait_above_3(6)], 
...             [47.902151605742517, 60.603197182023813, 72.965413990538366, 89.281859197449506])
True
```


Tests for q5_3 (points: None)
```
>>> np.allclose([predict_wait(x) for x in [1.5, 2.5]], 
...             [51.077412999812836, 57.427935787953487])
True
```
```
>>> np.allclose([predict_wait(x) for x in [3.5, 4.5]], 
...             [75.684821525023551, 81.123636593993922])
True
```
